# Analiza danych przestrzennych - ćwiczenia laboratoryjne 2023/2024

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania i oceny notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Jeżeli zestaw zadań wymaga skorzystania z fragmentów kodu opracowanego w ramach wcześniejszych zestawów zadań należy je umieścić we wskazanej komórce notatnika.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu, odpowiedzi uzupełniane manualnie nie podlegają ocenie.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają na celu wyłapać podstawowe błędy. Przejście przez kod wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie i zostanie ocenione na maksymalną liczbę punktów.
- Zadania należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zadania, które powodują wyświetlenie komunikatu o błędzie przerywającym wykonywanie kodu nie podlegają ocenie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem). Rozmiar poziomy figur nie powinien przekraczać 20 cali.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł.
- Zakresy osi wykresów przedstawiających rozkłady punktów w przestrzeni powinny być docięte do granic obszaru, na którym został wygenerowany proces punktowy.
- Proporcje osi wykresów przedstawiających rozkłady punktów w przestrzeni powinny być dobrane tak, aby wykresy odzwierciedlały rzeczywisty rozkład punktów w przestrzeni. Nie należy osiągać tego efektu manipulując rozmiarem całej figury.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autorów rozwiązania (nazwa zespołu oraz imiona, nazwiska i numery indeksów członków zespołu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

Nazwa zespołu:
Członkowie:

---

## Zestaw zadań 7: Analiza danych przestrzennych z wykorzystaniem dedykowanych pakietów (część 1)

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
import geopandas as gpd
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

### Zadanie 1: Import i przygotowanie danych (8 pkt)

Wykorzystując funkcję `geopandas.read_file()` zaimportuj do notatnika dane z trzech załączonych do niego plików i przygotuj na ich podstawie trzy tablice zawierające wybrane dane dotyczące województw, miast oraz dwóch rzek - Wisły i Odry:

1) Tablica `wojewodztwa` zawierająca następujące kolumny:
    - "Nazwa" - nazwy województw
    - "geometry" - geometria obiektów
<br/>
<br/>
2) Tablica `miasta` zawierająca następujące kolumny:
    - "Nazwa" - nazwy miast
    - "geometry" - geometria obiektów
<br/>
<br/>
3) Tablica `rzeki` zawierająca następujące kolumny:
    - "Nazwa" - nazwy rzek
    - "geometry" - geometria obiektów

Wyświetl przygotowane tablice.

Uwagi do wykonania zadania:
 - Pliki zawierają więcej danych niż opisane wyżej tablice. W wynikowych tablicach nie powinny znaleźć się zbędne dane.
 - Poprawny import polskich nazw może wymagać ustawienia odpowiedniej wartości argumentu `encoding` (utf-8 lub 1250) podczas wywoływania funkcji `geopandas.read_file()`.

<!-- grade_id: cell-9c5827e1303cedf5, points: 0, grade: false, solution: true -->
```python
df_woj=gpd.read_file("Wojewodztwa.zip")
df_woj
wojewodztwa=gpd.GeoDataFrame({'Nazwa':df_woj["JPT_NAZWA_"],'geometry':df_woj["geometry"]})
print(wojewodztwa)
```

Output:
```
Nazwa                                           geometry
0               śląskie  MULTIPOLYGON (((18.91685 51.09608, 18.91677 51...
1              opolskie  POLYGON ((17.81203 51.18669, 17.81209 51.18670...
2        świętokrzyskie  POLYGON ((19.70430 50.75239, 19.70443 50.75248...
3             pomorskie  POLYGON ((17.66483 54.78297, 17.66653 54.78334...
4             podlaskie  POLYGON ((21.98108 52.97859, 21.98078 52.97861...
5    zachodniopomorskie  POLYGON ((14.81157 54.04037, 14.89525 54.05570...
6          dolnośląskie  POLYGON ((15.10534 51.42730, 15.10540 51.42732...
7         wielkopolskie  POLYGON ((16.71447 53.29913, 16.71447 53.29920...
8          podkarpackie  POLYGON ((22.03538 50.80675, 22.03550 50.80681...
9           małopolskie  POLYGON ((19.97230 50.51625, 19.97227 50.51633...
10  warmińsko-mazurskie  POLYGON ((21.55932 54.32250, 21.55939 54.32250...
11              łódzkie  POLYGON ((19.99385 51.18395, 19.99410 51.18395...
12          mazowieckie  POLYGON ((21.98576
```

<!-- grade_id: cell-749fa2b3c9af223d, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert type(wojewodztwa) == gpd.geodataframe.GeoDataFrame
assert wojewodztwa.shape == (16, 2)
assert list(wojewodztwa.columns) == ['Nazwa', 'geometry']
```

<!-- grade_id: cell-2f6e5df2fa030235, points: 0, grade: false, solution: true -->
```python
df_miej=gpd.read_file("Miejscowosci.zip")
df_miej
miasta= df_miej[df_miej["rodzaj"]=='miasto']
miasta = gpd.GeoDataFrame({'Nazwa':miasta["nazwaGlown"],'geometry':miasta["geometry"]})
```

<!-- grade_id: cell-e319ae24aca025a4, points: 3, grade: true, solution: false -->
```python
# Komórka testowa
assert type(miasta) == gpd.geodataframe.GeoDataFrame
assert miasta.shape == (964, 2)
assert list(miasta.columns) == ['Nazwa', 'geometry']
```

<!-- grade_id: cell-e02cfe1cd85e0837, points: 0, grade: false, solution: true -->
```python
# Rzeki
 
```

Output:
```
Nazwa                                           geometry
0   Odra  MULTILINESTRING ((17.52056 49.61373, 17.52159 ...
1  Wisła  MULTILINESTRING ((19.00368 49.60595, 19.00338 ...
```

<!-- grade_id: cell-672b9a190228f2cd, points: 3, grade: true, solution: false -->
```python
# Komórka testowa
assert type(rzeki) == gpd.geodataframe.GeoDataFrame
assert rzeki.shape == (2, 2)
assert list(rzeki.columns) == ['Nazwa', 'geometry']
```

### Zadanie 2: Transformacja między odwzorowaniami kartograficznymi (2 pkt)

Wczytane dane są zapisane z wykorzystaniem różnych odwzorowań kartograficznych. Przed rozpoczęciem analizy należy je ze sobą uzgodnić.

Korzystając z własności obiektów GeoDataFrame `geopandas.GeoDataFrame.crs` sprawdź odwzorowanie kartograficzne, w którym zapisane są poszczególne tablice. Jeżeli odwzorowanie jest inne niż epsg:2180 dokonaj jego konwersji z wykorzystaniem funkcji `geopandas.GeoDataFrame.to_crs()`.

<!-- grade_id: cell-adfcf2bd09c20ba2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[9], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-9191c6595dac133c, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert wojewodztwa.crs == "epsg:2180"
assert miasta.crs == "epsg:2180"
assert rzeki.crs == "epsg:2180"
```

### Zadanie 3: Wizualizacja danych wejściowych (3 pkt)

Korzystając z funkcji pakietu matplotlib oraz funkcji `geopandas.GeoDataFrame.plot()` zwizualizuj komplet przygotowanych danych. Ustaw parametry wykresu oraz wizualizowanych obiektów (wielkość figury, kolory, grubość linii, wielkość punktów itp.) tak, aby wygenerowana mapa była przejrzysta i czytelna.

<!-- grade_id: cell-e4b44f2f3cc8497c, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Zadanie 4: Mierzenie długości, pola powierzchni i odległości (10 pkt)

#### a) Długość Wisły
Korzystając z własności `geopandas.GeoSeries.length` zmierz długość Wisły w metrach oraz:
 - zapisz wynik do zmiennej `length`,
 - wyświetl odpowiedź (pełnym zdaniem) na pytanie o długość Wisły podając odpowiedź w kilometrach.

<!-- grade_id: cell-15ff549b14ca0466, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-24758c9f3da5ccde, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert np.isclose(length, 1019757.7991167122)
```

#### b) Pola powierzchni województw
Korzystając z własności `geopandas.GeoSeries.area` zmierz powierzchnię województw w metrach kwadratowych oraz:
 - zapisz wyniki do tablicy `wojewodztwa` w kolumnie o nazwie "Pole", 
 - oblicz pole powierzchni Polski w metrach kwadartowych i zapisz wynik do zmiennej `area`,
 - wyświetl odpowiedź (pełnym zdaniem) na pytanie o powierzchnię Polski podając odpowiedź w kilometrach kwadratowych,
 - przygotuj wizualiację, na której kolor przypisany województwom jest powiązany z ich powierzchnią.

Wizualizacja powinna zawierać:
- województwa, wypełnione kolorem powiązanym z ich powierzchnią,

Ustaw parametry wykresu oraz wizualizowanych obiektów (wielkość figury, kolory, grubość linii, wielkość punktów itp.) tak aby wygenerowana mapa była przejrzysta i czytelna.

<!-- grade_id: cell-eb301c3344f17b1e, points: 0, grade: false, solution: true -->
```python
# Obliczenia
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-54e8fb04359f397e, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert np.all(np.isclose(wojewodztwa["Pole"].values, np.array([1.23170381e+10, 9.39998646e+09, 1.16972569e+10, 1.83069651e+10,
                                                               2.01931654e+10, 2.29033672e+10, 1.99361725e+10, 2.97972921e+10,
                                                               1.78437052e+10, 1.51660383e+10, 2.41505782e+10, 1.81944352e+10,
                                                               3.55292977e+10, 1.79475269e+10, 2.51340028e+10, 1.39898166e+10])))
assert np.isclose(area, 312506644577.3843)
```

<!-- grade_id: cell-ab1428f12655827b, points: 2, grade: true, solution: true -->
```python
# Wizualizacja
# YOUR CODE HERE
raise NotImplementedError()
```

#### c) Odległość miast od Krakowa
Korzystając z funkcji `geopandas.GeoSeries.distance()` zmierz odległość poszczególnych miast od Krakowa w metrach oraz:
 - zapisz wyniki do tablicy `miasta` w kolumnie o nazwie "Odległość od Krakowa", 
 - oblicz w jakiej odległości od Krakowa w kilometrach leży Poznań i zapisz wynik do zmiennej `distance`,
 - odpowiedz na pytanie (pełnym zdaniem) w jakiej odległości wyrażonej w kilometrach od Krakowa leży Poznań,
 - przygotuj wizualiację, na której kolor przypisany miastom jest powiązany z ich odległością od Krakowa. 

Wizualizacja powinna zawierać:
- kontury województw, bez wypełnienia,
- lokalicaje miast.

Ustaw parametry wykresu oraz wizualizowanych obiektów (wielkość figury, kolory, grubość linii, wielkość punktów itp.) tak aby wygenerowana mapa była przejrzysta i czytelna.

<!-- grade_id: cell-ce6d9808f4193a3c, points: 0, grade: false, solution: true -->
```python
# Obliczenia
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ea3373eed57cd3e3, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert np.isclose(miasta["Odległość od Krakowa"].min(), 0.0)
assert np.isclose(miasta["Odległość od Krakowa"].max(), 579020.726082644)
assert np.isclose(miasta["Odległość od Krakowa"].mean(), 271127.08178774477)
assert np.isclose(distance, 334.60495145299524)
```

<!-- grade_id: cell-8d841cc379d33cb6, points: 2, grade: true, solution: true -->
```python
# Wizualizacja
# YOUR CODE HERE
raise NotImplementedError()
```

### Zadanie 5: Wykorzystanie prostych operacji geometrycznych (12 pkt)

#### a) Miasta blisko Wisły

Korzystając z funkcji `geopandas.GeoSeries.buffer()` oraz `geopandas.GeoSeries.within()` sprawdź, które miast leżą bliżej niż 20 km od Wisły oraz:
 - zapisz wyniki do tablicy `miasta` w kolumnie o nazwie "Blisko Wisły" (True/False),
 - oblicz ile miast leży bliżej niż 20 km od Wisły i zapisz wynik do zmiennej `cities`,
 - odpowiedz na pytanie (pełnym zdaniem) ile miast leży bliżej niż 20 km od Wisły,
 - przygotuj wizualiację, na której oznaczony zostanie obszar znajdujący się bliżej niż 20 km od Wisły, a znajdujące się w tym obszarze miasta będą wyróżnione kolorem względem pozostałych miast.

Wizualizacja powinna zawierać:
- województwa, wypełnione jednolitym kolorem,
- miasta, w jednym z dwóch wybranych kolorów w zależności od spełnienia analizowanego kryterium.
- rzeki wypełnione jednolitym kolorem,
- oznaczenie terenu leżącego nie dalej niż 20 km od Wisły.

Ustaw parametry wykresu oraz wizualizowanych obiektów (wielkość figury, kolory, grubość linii, wielkość punktów itp.) tak aby wygenerowana mapa była przejrzysta i czytelna.

<!-- grade_id: cell-02e093a857861332, points: 0, grade: false, solution: true -->
```python
# Obliczenia
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b1ad6f1be0de19b9, points: 3, grade: true, solution: false -->
```python
# Komórka testowa
assert cities == 148
```

<!-- grade_id: cell-1d944ce0b7080e86, points: 3, grade: true, solution: true -->
```python
# Wizualizacja
# YOUR CODE HERE
raise NotImplementedError()
```

#### b) Województwa, przez które przepływa Odra
Korzystając z funkcji `geopandas.GeoSeries.intersects()` sprawdź, przez które województwa przepływa Odra oraz:
 - zapisz wyniki do tablicy `wojewodztwa` w kolumnie o nazwie "Odra w obrębie" (True/False) oraz do zmiennej `voivodeships` (lista województw, przez które przepływa Odra),
 - odpowiedz na pytanie (pełym zdaniem) przez które województa przepływa Odra,
 - przygotuj wizualiację, na której województwa przez które przepływa Odra będą wyróżnione kolorem względem pozostałych województw.

Wizualizacja powinna zawierać:
- województwa, wypełnione jednym z dwóch wybranych kolorów w zależności od spełnienia analizowanego kryterium,
- miasta, wypełnione jednolitym kolorem,
- rzeki wypełnione jednolitym kolorem.

Ustaw parametry wykresu oraz wizualizowanych obiektów (wielkość figury, kolory, grubość linii, wielkość punktów itp.) tak aby wygenerowana mapa była przejrzysta i czytelna.

<!-- grade_id: cell-07caea011ad0dfd6, points: 0, grade: false, solution: true -->
```python
# Obliczenia
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0e6fb598a6c4608f, points: 3, grade: true, solution: false -->
```python
# Komórka testowa
assert sorted(voivodeships) == ['dolnośląskie', 'lubuskie', 'opolskie', 'zachodniopomorskie', 'śląskie']
```

<!-- grade_id: cell-bc1504ffa6bf53aa, points: 3, grade: true, solution: true -->
```python
# Wizualizacja
# YOUR CODE HERE
raise NotImplementedError()
```